```python
import pandas as pd
import nose.tools
# Write your imports here
```

# Data Tidying and Cleaning Lab
## Reading, tidying and cleaning data. Preparing data for exploration, mining, analysis and learning

### Problem 1. Read the dataset (2 points)
The dataset [here](http://archive.ics.uci.edu/ml/datasets/Auto+MPG) contains information about fuel consumption in different cars.

Click the "Data Folder" link and read `auto_mpg.data` into Python. You can download it, if you wish, or you can read it directly from the link.

Give meaningful (and "Pythonic") column names, as per the `auto_mpg.names` file:
1. mpg
2. cylinders
3. displacement
4. horsepower
5. weight
6. acceleration
7. model_year
8. origin
9. car_name

<!-- grade_id: read_data, points: 0, grade: false, solution: true -->
```python
mpg_data = pd.read_fwf("../Data Tidying and Cleaning/auto-mpg.data", header = None)
```

```python
mpg_data.columns = ["mpg","cylinders","displacement","horsepower","weight","acceleration","model_year","origin","car_name"]
```

<!-- grade_id: read_data_tests, points: 2, grade: true, solution: false -->
```python
nose.tools.assert_is_not_none(mpg_data)
```

Print the first 4 rows in the dataset to get a feel of what it looks like:

<!-- grade_id: cell-80f1e6004aaafef8, points: 0, grade: false, solution: true -->
```python
mpg_data.head(4)
```

Output:
```
mpg  cylinders  displacement horsepower  weight  acceleration  model_year  \
0  18.0          8         307.0      130.0  3504.0          12.0          70   
1  15.0          8         350.0      165.0  3693.0          11.5          70   
2  18.0          8         318.0      150.0  3436.0          11.0          70   
3  16.0          8         304.0      150.0  3433.0          12.0          70   

   origin                     car_name  
0       1  "chevrolet chevelle malibu"  
1       1          "buick skylark 320"  
2       1         "plymouth satellite"  
3       1              "amc rebel sst"
```

### Problem 2. Inspect the dataset (1 point)
Write a function which accepts a dataset and returns the number of observations and features in it, like so: 

``` 10 observations on 15 features```

Where 10 and 15 should be replaced with the real numbers. Test your function with the `auto_mpg` dataset.

Make sure the function works with other datasets (don't worry about "1 features" or "1 observations", just leave it as it is).

<!-- grade_id: get_shape_function, points: 0, grade: false, solution: true -->
```python
def observations_and_features(dataset):
    """
    Returns the number of observations and features in the provided dataset
    """
    observations = len(mpg_data.axes[0])
    features = len(mpg_data.axes[1])
    return "{} observations on {} features".format(observations, features)
```

<!-- grade_id: get_shape_function_tests, points: 1, grade: true, solution: false -->
```python
print(observations_and_features(mpg_data))
```

Output:
```
398 observations on 9 features
```

Inspect the data types for each column.

<!-- grade_id: cell-152f652655c53f2a, points: 0, grade: false, solution: true -->
```python
mpg_data.dtypes
```

Output:
```
mpg             float64
cylinders         int64
displacement    float64
horsepower       object
weight          float64
acceleration    float64
model_year        int64
origin            int64
car_name         object
dtype: object
```

### Problem 3. Correct errors (1 point)
The `horsepower` column looks strange. It's a string but it must be a floating-point number. Find out why this is so and convert it to floating-point number.

<!-- grade_id: convert_to_numeric, points: 0, grade: false, solution: true -->
```python
mpg_data['horsepower'] = pd.to_numeric(mpg_data['horsepower'],errors='coerce')
```

<!-- grade_id: convert_to_numeric_tests, points: 1, grade: true, solution: false -->
```python
nose.tools.assert_equal(mpg_data.horsepower.dtype, "float64")
```

### Problem 4. Missing values: inspection (1 point)
We saw that the `horsepower` column contained null values. Display the rows which contain those values. Assign the resulting dataframe to the `unknown_hp` variable.

<!-- grade_id: unknown_hp, points: 0, grade: false, solution: true -->
```python
def get_unknown_hp(dataframe): 
    unknown_hp = mpg_data[mpg_data.horsepower.isnull()]
    unknown_hp = unknown_hp.horsepower
    return unknown_hp
```

<!-- grade_id: unknown_hp_tests, points: 1, grade: true, solution: false -->
```python
cars_with_unknown_hp = get_unknown_hp(mpg_data)
print(cars_with_unknown_hp)
```

Output:
```
32    NaN
126   NaN
330   NaN
336   NaN
354   NaN
374   NaN
Name: horsepower, dtype: float64
```

### Problem 5. Missing data: correction (1 point)
It seems like the `NaN` values are a small fraction of all values. We can try one of several things:
* Remove them
* Replace them (e.g. with the mean power of all cars)
* Look up the models on the internet and try our best guess on the power

The third one is probably the best but the first one will suffice since these records are too few. Remove those values. Save the dataset in the same `mpg_data` variable. Ensure there are no more `NaN`s.

<!-- grade_id: remove_nulls, points: 0, grade: false, solution: true -->
```python
mpg_data = mpg_data.dropna()
mpg_data.isnull().values.any()
```

Output:
```
False
```

<!-- grade_id: remove_nulls_test, points: 1, grade: true, solution: false -->
```python
nose.tools.assert_equal(len(get_unknown_hp(mpg_data)), 0)
```

### Problem 6. Years of production (1 + 1 points)
Display all unique model years. Assign them to the variable `model_years`.

<!-- grade_id: model_years, points: 0, grade: false, solution: true -->
```python
def get_unique_model_years(dataframe):
    model_years = mpg_data['model_year'].unique()
    return model_years
```

<!-- grade_id: model_years_test, points: 1, grade: true, solution: false -->
```python
model_years = get_unique_model_years(mpg_data)
print(model_years)
```

Output:
```
[70 71 72 73 74 75 76 77 78 79 80 81 82]
```

These don't look so good. Convert them to real years, like `70 -> 1970, 71 -> 1971`. Replace the column values in the dataframe.

## We use .map to call the function for every element in the array instead of adding 1900 to the number , because every time we start the code it will add 1900 to the new year ---> 1970 ---> 3870 ...


<!-- grade_id: model_year, points: 0, grade: false, solution: true -->
```python
mpg_data.model_year = mpg_data.model_year.map({
    70: 1970,
    71: 1971,
    72: 1972,
    73: 1973,
    74: 1974,
    75: 1975,
    76: 1976,
    77: 1977,
    78: 1978,
    79: 1979,
    80: 1980,
    81: 1981,
    82: 1982
})
```

<!-- grade_id: model_year_test, points: 1, grade: true, solution: false -->
```python
model_years = get_unique_model_years(mpg_data)
print(model_years)
```

Output:
```
[nan]
```

### Problem 7. Exploration: low-power cars (1 point)
The data looks quite good now. Let's try some exploration.

Write a function to find the cars which have the smallest number of cylinders and print their model names. Return a list of car names.

<!-- grade_id: car_names, points: 0, grade: false, solution: true -->
```python
def get_model_names_smallest_cylinders(dataframe):
    car_names = mpg_data.loc[mpg_data.cylinders <= mpg_data.cylinders.min()].car_name
    return car_names
```

<!-- grade_id: car_names_test, points: 1, grade: true, solution: false -->
```python
car_names = get_model_names_smallest_cylinders(mpg_data)
print(car_names)
nose.tools.assert_true(car_names.shape == (4,) or car_names.shape == (4, 1))
```

Output:
```
71     "mazda rx2 coupe"
111          "maxda rx3"
243         "mazda rx-4"
334      "mazda rx-7 gs"
Name: car_name, dtype: object
```

### Problem 8. Exploration: correlations (1 point)
Finally, let's see some connections between variables. These are also called **correlations**.

Find how to calculate correlations between different columns using `pandas`.

**Hint:** The correlation function in `pandas` returns a `DataFrame` by default. You need only one value from it.

Create a function which accepts a dataframe and two columns and prints the correlation coefficient between those two columns.

<!-- grade_id: correlation, points: 0, grade: false, solution: true -->
```python
def calculate_correlation(dataframe, first_column, second_column):
    correlation = mpg_data[first_column].corr(mpg_data[second_column])
    return correlation
```

<!-- grade_id: cell-457c5946f2350991, points: 1, grade: true, solution: false -->
```python
hp_weight = calculate_correlation(mpg_data, "horsepower", "weight")
print("Horsepower:Weight correlation coefficient:", hp_weight)
nose.tools.assert_almost_equal(hp_weight, 0.864537737574, delta = 0.01)

```

Output:
```
Horsepower:Weight correlation coefficient: 0.8645377375741444
```

```python

```

```python

```